--- 
title_meta  : Chapter 4
title       : Application to Expression Data
description : "In this chapter, you will combine what you learned in the first three chapters to read in, tidy, and visualize gene expression data. Let's get started!"

--- type:NormalExercise xp:100 skills:1
## Putting it all together

*** =hint
There is no code to type, just hit Submit Answer.

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
```

*** =solution
```{r}
```

*** =sct
```{r}
success_msg("Now that you understand the study, let's analyze some data!")
```

